Take a few minutes for each problem. Skip a problem if you get stuck because you can always go back to it later. 

Make sure you fill in any place that says `YOUR CODE HERE` with your code. Also, don't forget to remove the `pass` lines as well.

# DESN210/BTEC420 Quiz 0

## Instructions

The quiz is **open book**. You can use any class notes that you may have as well as any of the material that you can download from [Canvas](https://canvas.case.edu/). You are **not** able to get help from other students for this quiz.

The name of the notebook file, `Quiz0.ipynb`, will need to remain the same in order for me to properly download and grade it.

## Datatypes (6 points)

You know your account balance and you go out to make a purchase. Update `calculate_new_balance` to return the new account balance.

<!-- grade_id: cell-ecde1ba5d9a569d9, points: 0, grade: false, solution: true -->
```python
def calculate_new_balance(account_balance, purchase):
    # YOUR CODE HERE
    return account_balance-purchase
```

<!-- grade_id: cell-0cd8e34deb86715a, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
observed = calculate_new_balance(45098, 12075)
expected = 33023
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = calculate_new_balance(10, 11)
expected = -1
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

We know a user's first and last name. Update `full_name` to return a string that is equal to the first name and last name joined together with a space in the middle.

<!-- grade_id: cell-2286feff08917d07, points: 0, grade: false, solution: true -->
```python
def full_name(first_name, last_name):
    # YOUR CODE HERE
    return first_name+" "+last_name
```

<!-- grade_id: cell-c896c95cfaca10a1, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
observed = full_name("John", "Doe")
expected = "John Doe"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = full_name("Jane", "Smith")
expected = "Jane Smith"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

Update `cash` to return a decimal object with a value of 1.33.

<!-- grade_id: cell-161605f13e14bdc6, points: 0, grade: false, solution: true -->
```python
from decimal import Decimal

def cash():
    # YOUR CODE HERE
    return Decimal.from_float(1.33)
```

<!-- grade_id: cell-143511ef840610e9, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
value = cash()
assert isinstance(value, Decimal)
assert value.as_tuple().digits == (1, 3, 3)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-19-d5bb3148d46e> in <module>
      4 value = cash()
      5 assert isinstance(value, Decimal)
----> 6 assert value.as_tuple().digits == (1, 3, 3)

AssertionError:
```

Return the first character of the provided string `name`.

<!-- grade_id: cell-41b5ca12e69ee2c6, points: 0, grade: false, solution: true -->
```python
def first_character(name):
    # YOUR CODE HERE
    return list(name)[0]
```

<!-- grade_id: cell-ff5c751b602068b2, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
observed = first_character("Milo")
expected = "M"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = first_character("Bill Gates")
expected = "B"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

Return the last character of the provided string `name`.

<!-- grade_id: cell-9e2d6a0c1391150c, points: 0, grade: false, solution: true -->
```python
def last_character(name):
    # YOUR CODE HERE
    return list(name)[-1]
```

<!-- grade_id: cell-22abaaba3c487331, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
observed = last_character("Milo")
expected = "o"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = last_character("Bill Gates")
expected = "s"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

Return the length of the provided string `name`. Use the builtin [len](https://docs.python.org/3/library/functions.html#len) function.

<!-- grade_id: cell-b7b671dcf50e7ce5, points: 0, grade: false, solution: true -->
```python
def string_len(name):
    # YOUR CODE HERE
    return len(name)
```

<!-- grade_id: cell-19089c7a9184fbd5, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
observed = string_len("Milo")
expected = 4
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = string_len("Bill Gates")
expected = 10
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

## Lists (6 points)

Return a list of integers. Any integers will do as long as there are more than one in the list.

<!-- grade_id: cell-47bd2b043d78f2ac, points: 0, grade: false, solution: true -->
```python
def numbers():
    # YOUR CODE HERE
    from random import randint
    numbers=[]
    a=randint(1,10)
    for i in range(a):
        numbers.append(randint(1,10))
    return numbers
```

<!-- grade_id: cell-83d7c32048d6a491, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
rv = numbers()
assert len(rv) > 1, "Does your list have more than 1 number?"
assert all(isinstance(n, int) for n in rv), \
    "Does your list contain only integers?"
```

Use the correct index to pull data from the list. Update `joker` and use the indexing capability of lists to return the last value in the list.

<!-- grade_id: cell-f86f0a4ffaf52d12, points: 0, grade: false, solution: true -->
```python
cards = ["Ace", "King", "Queen", "Joker"]

def joker(cards):
    # YOUR CODE HERE
    return cards[cards.index("Joker")]
   
```

<!-- grade_id: cell-010f581e2a503d0d, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
assert joker(cards) == "Joker"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

Using a slice, return everything except the first and last elements of the provided list.

<!-- grade_id: cell-5e5f28a7dca795e9, points: 0, grade: false, solution: true -->
```python
def middle(items):
    # YOUR CODE HERE
    return items[1:-1]
```

<!-- grade_id: cell-6b165fc3a3a003df, points: 2, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
assert middle(["Ace", "King", "Queen", "Joker"]) == ["King", "Queen"]
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

assert middle(list(range(8))) == list(range(1, 7))
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

Append the new value from `val` to the end of the `things` list.

<!-- grade_id: cell-0e5f8d9f334f8631, points: 0, grade: false, solution: true -->
```python
def add_to_things(val, things):
    # YOUR CODE HERE
    things.append(val)
    return things
```

<!-- grade_id: cell-03b476068965e448, points: 2, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
t0 = []
add_to_things("apple", t0)
assert t0 == ["apple"]
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

t1 = [""]
add_to_things("laptop", t1)
assert t1 == ["", "laptop"]
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

## Dictionaries (5 points)

Return an empty dictionary.

<!-- grade_id: cell-1be66c3b435f6bad, points: 0, grade: false, solution: true -->
```python
def empty_dict():
    # YOUR CODE HERE
    return {}
```

<!-- grade_id: cell-8ff1171078b95e5f, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
rv = empty_dict()
assert isinstance(rv, dict), "You have not returned a dictionary."
assert len(rv) == 0, "Your dictionary is not empty."
```

Return a dictionary with a single key called `name` and the value should be your first name.

<!-- grade_id: cell-5de98b9e5525e699, points: 0, grade: false, solution: true -->
```python
def me():
    # YOUR CODE HERE
    return {'name':"Weihao"}
```

<!-- grade_id: cell-70ff84d25367dcfd, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
rv = me()
assert isinstance(rv, dict), "You have not returned a dictionary."
assert rv.get("name"), "Your dictionary does not have a name key."
assert isinstance(rv["name"], str), \
    "It doesn't look like you dictionary contains your name."
```

Add a new key to a dictionary. The dictionary is named `things`, the key is named `key` and the value is named `value`.

<!-- grade_id: cell-ca8308d54b7710b1, points: 0, grade: false, solution: true -->
```python
def add_key_value(things, key, value):
    # YOUR CODE HERE
    things[key]=value
```

<!-- grade_id: cell-6fb3584e5fd8203b, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
t0 = {}
add_key_value(t0, "fav", "pizza")
expected = {"fav": "pizza"}
assert t0 == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

t1 = {"name": "John"}
add_key_value(t1, "address", "123 Main St.")
expected = {"name": "John", "address": "123 Main St."}
assert t1 == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

Return a value for `key` from the provided dictionary `data`. If the key doesn't exist in the dictionary the return the string `"not found"`.

<!-- grade_id: cell-6ec2f6a6112c49ef, points: 0, grade: false, solution: true -->
```python
def safe_get(data, key):
    # YOUR CODE HERE
    try:
        return  data[key]
    except:
        return "not found"
    
```

<!-- grade_id: cell-47e869c57a7a6849, points: 2, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
observed = safe_get({}, "anything")
expected = "not found"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = safe_get({"watch": "tv"}, "watch")
expected = "tv"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

### Conditionals (3 points)

Use an `if-elif-else` to implement the following logic given two names `name0` and `name1`:

 1. If `name0` and `name1` are the same length return `"same"`
 2. If `name0` is longer than `name1` return `"longer"`
 3. If `name0` is shorter than `name1` return `"shorter"`
 
 ***Hint:*** Use the `len()` builtin that you used earlier!

<!-- grade_id: cell-869277c9bd9fd9c0, points: 0, grade: false, solution: true -->
```python
def compare_names(name0, name1):
    # YOUR CODE HERE
    if len(name0) == len(name1):
        return "same"
    elif len(name0) > len(name1):
        return "longer"
    else:
        return "shorter"
```

<!-- grade_id: cell-008b0646f8a9f891, points: 3, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#
observed = compare_names("Dodge", "Chevy")
expected = "same"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = compare_names("Dodge", "BMW")
expected = "longer"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = compare_names("Audi", "Mercedes")
expected = "shorter"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```